---
title: 'Changing the inputs'
description: 'The learner will be able to use Google sheets to add features to a binomial model that will make it customizable, i.e. the user will be able to quickly and easily change inputs, which will automatically lead to changes in the model.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
